Sir2 and Fun30 regulate ribosomal DNA replication timing via MCM helicase positioning and nucleosome occupancy

Author response:

The following is the authors’ response to the previous reviews.

eLife Assessment

This valuable study is a detailed investigation of how chromatin structure influences replication origin function in yeast ribosomal DNA, with a focus on the role of the histone deacetylase Sir2 and the chromatin remodeler Fun30. Convincing evidence shows that Sir2 does not affect origin licensing but rather affects local transcription and nucleosome positioning which correlates with increased origin firing. Overall, the evidence is solid and the model plausible. However, the methods employed do not rigorously establish a key aspect of the mechanism where initiation precisely occurs or rigorously exclude alternative models and the effect of Sir2 on transcription is not re-examined in the fun30 context.

Clarification on Sir2 Effect on Transcription in the fun30 Context

We appreciate the reviewers’ thorough assessment but would like to clarify that the effect of Sir2 on transcription in the fun30 context was addressed in both the original and revised manuscripts. However, we recognize that the presentation of the qPCR results may have been unclear, as we initially plotted absolute transcript levels without normalizing for rDNA array size differences among the genotypes. We have now corrected this.

After normalizing for copy number variations, the qPCR data show that the sir2 fun30 double mutant results in a ~40-fold increase in C-pro transcription relative to WT, compared to a 4-fold and 19-fold increase in fun30 and sir2 single mutants, respectively (Figure 5, figure supplement 6). These results have been discussed in the manuscript result section, where we note that "C-pro RNA levels were approximately twice as high in sir2 fun30 compared to sir2 cells when adjusted for rDNA size differences." This observation is critical for addressing both alternative models of MCM disappearance and for pinpointing transcription initiation sites, as detailed in the following sections.

Public Reviews:

Reviewer #1 (Public review):

Summary:

This paper presents a mechanistic study of rDNA origin regulation in yeast by SIR2. Each of the ~180 tandemly repeated rDNA gene copies contains a potential replication origin. Earlyefficient initiation of these origins is suppressed by Sir2, reducing competition with origins distributed throughout the genome for rate-limiting initiation factors. Previous studies by these authors showed that SIR2 deletion advances replication timing of rDNA origins by a complex mechanism of transcriptional de-repression of a local PolII promoter causing licensed origin proteins (MCMcomplexes) to re-localize (slide along the DNA) to a different (and altered) chromatin environment. In this study, they identify a chromatin remodeler, FUN30, that suppresses the sir2∆ effect, and remarkably, results in a contraction of the rDNA to about onequarter it's normal length/number of repeats, implicating replication defects of the rDNA. Through examination of replication timing, MCM occupancy and nucleosome occupancy on the chromatin in sir2, fun30, and double mutants, they propose a model where nucleosome position relative to the licensed origin (MCM complexes) intrinsically determines origin timing/efficiency. While their interpretations of the data are largely reasonable and can be interpreted to support their model, a key weakness is the connection between Mcm ChEC signal disappearance and origin firing. While the cyclical chromatin association-dissociation of MCM proteins with potential origin sequences may be generally interpreted as licensing followed by firing, dissociation may also result from passive replication and as shown here, displacement by transcription and/or chromatin remodeling. Moreover, linking its disappearance from chromatin in the ChEC method with such precise resolution needs to be validated against an independent method to determine the initiation site(s). Differences in rDNA copy number and relative transcription levels also are not directly accounted for, obscuring a clearer interpretation of the results. Nevertheless, this paper makes a valuable advance with the finding of Fun30 involvement, which substantially reduces rDNA repeat number in sir2∆ background. The model they develop is compelling and I am inclined to agree, but I think the evidence on this specific point is purely correlative and a better method is needed to address the initiation site question. The authors deserve credit for their efforts to elucidate our obscure understanding of the intricacies of chromatin regulation. At a minimum, I suggest their conclusions on these points of concern should be softened and caveats discussed. Statistical analysis is lacking for some claims.

Strengths are the identification of FUN30 as suppressor, examination of specific mutants of FUN30 to distinguish likely functional involvement. Use of multiple methods to analyze replication and protein occupancies on chromatin. Development of a coherent model.

Weaknesses are failure to address copy number as a variable; insufficient validation of ChEC method relationship to exact initiation locus; lack of statistical analysis in some cases.

Review of revised version and response letter:

In the response, the authors make some improvements by better quantifying 2D gels, adding some missing statistical analyses, analyzing the effect of fun30 on rDNA replication in strains with reduced rDNA copy number, and using ChIP-seq of MCMs to support the ChEC-seq data. However, these additions do not address the main issue that is at the heart of their model: where initiation precisely occurs and whether the location is altered in the mutant(s). Thus, mechanistic insight is limited.

We discuss the issue regarding the initiation site below.

Under the section "Addressing Alternative Explanations", the authors claim that processes like transcription and passive replication cannot affect the displaced complex specifically. Why? They are not on same DNA (as mentioned in the Fig 1 legend).

Premature origin activation, not transcription, drives the disappearance of repositioned MCM complexes in sir2 mutants in HU.

Indeed, the reviewer is correct in suggesting that C-pro transcription confined to rDNA units with repositioned MCM complexes could selectively displace those complexes, potentially explaining the selective disappearance of displaced MCMs in sir2 cells. However, our analysis of C-pro transcription and MCM occupancy in G1 versus HU across the genotypes allows us to rule out this possibility.

We show that the fraction of repositioned MCMs in G1 cells is proportional to the level of C-pro transcription (WT < fun30 << sir2 < sir2 fun30), consistent with the involvement of transcription in the repositioning process during MCM loading in G1. Accordingly, with approximately twice the transcription in sir2 fun30 compared to sir2, we observe more repositioned MCMs in sir2 fun30 cells than in sir2 cells in G1 (Fig 5C).

However, if the disappearance of repositioned MCMs in HU were solely due to C-pro transcription rather than origin activation, we would expect the repositioned MCMs to disappear more quickly in sir2 fun30 cells. Contrary to this expectation, our data show that repositioned MCM complexes are more stable in sir2 fun30 mutants compared to sir2 mutants, indicating that transcription is not the primary factor in the disappearance of displaced MCM complexes in HU; rather, rDNA origin activation appears to be the key factor.

Replication initiation site in sir2. Using multiple independent approaches, including 2D gels, ChIP-seq, and EdU incorporation, we have demonstrated that rDNA origins fire prematurely in sir2 mutants, a conclusion that the reviewer does not contest. Once an origin fires, the MCM signal disappears from the site of its initial deposition, as expected, and this is confirmed in our MCM ChIP and HU ChEC data, both at rDNA origins and across the genome.

Given that the majority of MCM complexes in sir2 mutants are repositioned, it is expected that these repositioned complexes disappear following premature origin activation. With less than half of the licensed origins (or <30% of total rDNA copies) retaining MCM at non-repositioned sites in sir2 mutants, if only these non-repositioned complexes were firing, and the repositioned MCM complexes were disappearing via mechanisms other than replication initiation (e.g., transcription), rDNA replication in sir2 mutants would be severely compromised rather than accelerated. Given this, and the strong experimental evidence that repositioned MCM complexes fire prematurely, continued focus on alternative explanations for MCM complex disappearance seems unwarranted.

We present this analysis in the results section as follows:

“Finally, although deletion of FUN30 could suppress replication initiation at the rDNA either by inhibiting the firing of the active, repositioned MCM complex or by preventing MCM repositioning to the "active location" in the first place, our results suggest that suppression occurs through the former mechanism. Consistent with previous reports that fun30 mutants are deficient in transcriptional silencing (Neves-Costa et al. 2009), C-pro RNA levels were approximately twice as high in sir2 fun30 cells compared to sir2 cells when adjusted for rDNA size (Figure 5—figure supplement 6).

Moreover, deletion of FUN30 shifts the distribution toward the repositioned MCM location over the non-repositioned one in G1 cells (Figure 5C), aligning with the increased C-pro transcription observed in fun30 mutants. This shift is evident in both sir2 and SIR2 cells. Despite the increased transcription-mediated repositioning in sir2 fun30 cells compared to sir2 cells during G1, repositioned MCM persists longer in sir2 fun30 cells than in sir2 cells after release into HU. Additionally, sir2 fun30 mutants exhibit reduced MCM accumulation at the RFB compared to sir2 mutants after release into HU, supporting the conclusion that MCM disappearance in HU reflects origin activation rather than transcription-mediated displacement.”

The model in Fig 7 implies that initiation sites are different in WT versus the mutants and this determines their timing/efficiency. But they also suggest that the same site might be used with different efficiencies in this response. I agree that both are possibilities and are not resolved.

Adjustment of the model to account for repositioned MCMs in WT cells In Figure 5—figure supplement 5, we demonstrate that even in WT cells, a small fraction of repositioned MCMs (~5%) can be detected, and that these repositioned MCM complexes disappear prematurely. However, because this represents a very small fraction of MCMs in WT cells, we initially did not include it in our overall model in Figure 7. In light of the reviewer's comment, we have now revised the model to incorporate this detail.

Supporting their model requires better resolution to determine the actual replication initiation site. While this may be challenging, it should be feasible with methods to map nascent strands like DNAscent, or Okazaki fragment mapping.

The initiation site in sir2 mutants has been thoroughly analyzed and supported by extensive experimental data, as discussed above. While high-resolution techniques such as DNAscent or Okazaki fragment mapping could potentially offer another layer of validation, the likelihood of obtaining finer detail that would change the conclusions is minimal. The methods we employed provide sufficient resolution to pinpoint the initiation site, and our results align consistently with established replication models.

Further experimentation would not only be redundant but also unlikely to provide new insights beyond revalidation. Given the strength of our current data, we believe the conclusions regarding replication initiation are robust and well-supported, making additional experiments unnecessary at this stage. Our priority is to focus on advancing other aspects of the research that require deeper exploration.

The 2D gel analysis of strains with reduced rDNA copy numbers adequately addresses the copy number variable with regard to the replication effect.

Overall, the paper is improved by providing additional data and improved analysis. The paper nicely characterizes the effect of Fun30. The model is reasonable but remains lacking in precise details of mechanism.

Reviewer #2 (Public review):

Summary:

In this manuscript, the authors follow up on their previous work showing that in the absence of the Sir2 deacetylase the MCM replicative helicase at the rDNA spacer region is repositioned to a region of low nucleosome occupancy. Here they show that the repositioned displaced MCMs have increased firing propensity relative to non-displaced MCMs. In addition, they show that activation of the repositioned MCMs and low nucleosome occupancy in the adjacent region depend on the chromatin remodeling activity of Fun30.

Strengths:

The paper provides new information on the role of a conserved chromatin remodeling protein in regulation of origin firing and in addition provides evidence that not all loaded MCMs fire and that origin firing is regulated at a step downstream of MCM loading.

Weaknesses:

The relationship between the authors results and prior work on the role of Sir2 (and Fob1) in regulation of rDNA recombination and copy number maintenance is not explored, making it difficult to place the results in a broader context. Sir2 has previously been shown to be recruited by Fob1, which is also required for DSB formation and recombination-mediated changes in rDNA copy number. Are the changes that the authors observe specifically in fun30 sir2 cells related to this pathway? Is Fob1 required for the reduced rDNA copy number in fun30 sir2 double mutant cells?

Reviewer #3 (Public review):

Summary:

Heterochromatin is characterized by low transcription activity and late replication timing, both dependent on the NAD-dependent protein deacetylase Sir2, the founding member of the sirtuins. This manuscript addresses the mechanism by which Sir2 delays replication timing at the rDNA in budding yeast. Previous work from the same laboratory (Foss et al. PLoS Genetics 15, e1008138) showed that Sir2 represses transcription-dependent displacement of the Mcm helicase in the rDNA. In this manuscript, the authors show convincingly that the repositioned Mcms fire earlier and that this early firing partly depends on the ATPase activity of the nucleosome remodeler Fun30. Using read-depth analysis of sorted G1/S cells, fun30 was the only chromatin remodeler mutant that somewhat delayed replication timing in sir2 mutants, while nhp10, chd1, isw1, htl1, swr1, isw2, and irc5 had no effect. The conclusion was corroborated with orthogonal assays including two-dimensional gel electrophoresis and analysis of EdU incorporation at early origins. Using an insightful analysis with an Mcm-MNase fusion (Mcm-ChEC), the authors

show that the repositioned Mcms in sir2 mutants fire earlier than the Mcm at the normal position in wild type. This early firing at the repositioned Mcms is partially suppressed by Fun30. In addition, the authors show Fun30 affects nucleosome occupancy at the sites of the repositioned Mcm, providing a plausible mechanism for the effect of Fun30 on Mcm firing at that position. However, the results from the MNAse-seq and ChEC-seq assays are not fully congruent for the fun30 single mutant. Overall, the results support the conclusions providing a much better mechanistic understanding how Sir2 affects replication timing at rDNA,

Strengths

(1) The data clearly show that the repositioned Mcm helicase fires earlier than the Mcm in the wild type position.

(2) The study identifies a specific role for Fun30 in replication timing and an effect on nucleosome occupancy around the newly positioned Mcm helicase in sir2 cells.

Weaknesses

(1) It is unclear which strains were used in each experiment.

(2) The relevance of the fun30 phospho-site mutant (S20AS28A) is unclear.

(3) For some experiments (Figs. 3, 4, 6) it is unclear whether the data are reproducible and the differences significant. Information about the number of independent experiments and quantitation is lacking. This affects the interpretation, as fun30 seems to affect the +3 nucleosome much more than let on in the description.

Recommendations for the authors:

Reviewer #2 (Recommendations for the authors):

The authors have addressed my concerns by the addition of new experiments and analysis.

One point remains unclear regarding additional support for the Mcm-ChEC results using ChIP experiments to verify whether MCM redistributes in sir2D cells. In their rebuttal, the authors state that, "New supporting based evidence: ChIP at rDNA Origins. Our ChIP analysis also shows that the disappearance of the MCM signal at rDNA origins in sir2Δ cells released into HU is accompanied by signal accumulation at the replication fork barrier (RFB), indicative of stalled replication forks at this location (Figure 5 figure supplement 3)...." The ChIP data in Figure 5 supplement 3 show accumulation of the Mcm2 ChIP signal to the left of the RFB in sir2D cells but it doesn't look like there is any decrease in the MCM signal in sir2D relative to wild-type cells for the peak C-Pro. There is a new MCM peak suggesting perhaps a new MCM loading event.

Figure 5 figure supplement 3 shows the relative abundance of the MCM ChIP signal across the ~2 kb rDNA region, spanning from the MCM loading site at the rDNA origin (on the left) to the replication fork barrier (RFB) on the right. The MCM-ChIP data are normalized to the highest signal within this rDNA region rather than across the entire genome, meaning that only the relative abundance of MCM within this region is represented, and not comparisons between different conditions. We have now presented the results with the same axes for both alpha factor and HU.

In wild-type (WT) cells, the MCM signal remains primarily at the initial loading site. However, in sir2 mutants, a significant portion of the MCM signal shifts rightward, consistent with rDNA origin activation and the movement of MCM along with the progressing replication fork. While some replication forks stall at the RFB, others are positioned between the MCM loading site and the RFB. The additional MCM peak observed does not represent a new MCM loading event, as the experiment was conducted during S-phase, when new MCM loading is not possible.

Reviewer #3 (Recommendations for the authors):

In this revision the authors addressed my concerns and improved the manuscript and the presentation of the data. All my recommendations were implemented.

Changes from the earlier version to the revised version

Abstract

Abstract

The association between late replication timing and low transcription rates in eukaryotic heterochromatin is well-known, yet the specific mechanisms underlying this link remain uncertain. In Saccharomyces cerevisiae, the histone deacetylase Sir2 is required for both transcriptional silencing and late replication at the repetitive ribosomal DNA arrays (rDNA). We have previously reported that in the absence of SIR2, a derepressed RNA PolII repositions MCM replicative helicases from their loading site at the ribosomal origin, where they abut well-positioned, high-occupancy nucleosomes, to an adjacent region with lower nucleosome occupancy. By developing a method that can distinguish activation of closely spaced MCM complexes, here we show that the displaced MCMs at rDNA origins have increased firing propensity compared to the nondisplaced MCMs. Furthermore, we found that both, activation of the repositioned MCMs and low occupancy of the adjacent nucleosomes critically depend on the chromatin remodeling activity of FUN30. Our study elucidates the mechanism by which Sir2 delays replication timing, and it demonstrates, for the first time, that activation of a specific replication origin in vivo relies on the nucleosome context shaped by a single chromatin remodeler.

Introduction

SIR2 is the founding member of a family of NAD-dependent protein deacylases called sirtuins, which have important effects on transcription, replication, chromatin structure and metabolism in organisms from yeast to humans (Houtkooper et al. 2012; Shahgaldi and Kahmini 2021). SIR2 was initially identified in budding yeast in 1979 on the basis of its role in transcriptional repression at the silent mating type loci (Haber and George 1979; Klar et al. 1979; Rine et al. 1979), and it was therefore dubbed a "Silent Information Regulator". The silent mating type loci are heterochromatic regions that contain non-transcribed copies of the small open reading frames that determine yeast mating type, and deletion of SIR2 not only activates their transcription but also abolishes their heterochromatic structure (Ivy et al. 1986; Rine and Herskowitz 1987; Weiss and Simpson 1998; Ravindra et al. 1999; Haber 2012). It was later discovered that, in addition to regulating transcription at the silent mating type loci, Sir2 also regulates DNA replication, and it does so not only at the silent mating type loci but also at the two other heterochromatic regions of the yeast genome, namely at telomeres and the rDNA (Pasero et al. 2002; Yoshida et al. 2014). It is this last activity of Sir2, namely its effect on replication at the rDNA, that is the focus of this report.

The rDNA, which encodes the structural RNA components of the ribosome, is composed of approximately 150 tandemly arranged 9.1 kb repeats (Woolford and Baserga 2013). At almost 1.4 megabases, the rDNA is too large to be replicated passively, and thus it must contain its own origins of replication. This is achieved by the organization shown in Figure 1A: The ribosomal RNAs are produced as two major species, known as the 35S and 5S, which are transcribed by RNA polI and RNA polIII, respectively (Fernandez-Pevida et al. 2015). Between the 5’ ends of these genes is an origin of replication, and immediately to the left of this origin is a short Sir2-repressed RNA PolII-transcribed non-coding transcript called C-pro (Kobayashi and Ganley 2005; Li et al. 2006; Vasiljeva et al. 2008), which we will discuss further below. In wild-type cells, rDNA origins fire in the second half of S phase, making the rDNA among the later replicating regions of the genome, while in sir2 mutants, they fire early in S phase (Pasero et al. 2002; Yoshida et al. 2014; Foss et al. 2017).

rDNA structure, Chromatin Endogenous Cleavage (ChEC), and MCM displacement in sir2.

A. The 1.4 megabase rDNA region on chromosome XII is composed of approximately one hundred and fifty 9.1 kb repeats. Each repeat encodes both the 35S and 5S ribosomal RNAs (rRNA), which are transcribed by RNA Polymerase I (Pol I) and RNA Polymerase III (Pol III), respectively. The ribosomal origin of replication (rARS) is located between the 5’ ends of these genes. The C-pro transcript, which is suppressed by Sir2, initiates approximately two hundred base pairs from the rARS. The unidirectional replication fork block, which functions to prevent collision between replication and transcription machinery, is depicted in green. The probe used in Southern blots in Figure 3 to assess licensing is marked by *. B. In ChEC-seq, micrococcal nuclease (MNase; depicted as scissors) is fused to the protein of interest, in this case MCM2. The double-hexameric MCM helicase complex, MCM2-7, is depicted in purple, and nucleosomes are in blue. Cleavage is induced by addition of calcium to permeabilized cells. Due to the proximity of nucleosome entrance and exit sites, MCM2-ChEC can reveal not only the binding site of the MCM complex but also that of the adjacent nucleosome. C. Depiction of nucleosomes and MCM complex in wild type and sir2 in G1 and G2. Nucleosomes are numbered with respect to C-pro transcription. De-repression of C-pro transcription in sir2 causes RNA PolII to displace the MCM complex to the right. The three different MCM helicase complexes depicted in the bottom right panel are intended to convey the three most prominent locations for the complex in sir2; this is not intended to indicate that there is ever more than one MCM complex in the same rDNA repeat. Created with BioRender.com.

An indication of the biological significance of maintaining the late replication status of the rDNA came from a genome wide association study aimed at identifying naturally occurring genetic variants that affect replicative lifespan (Kwan et al. 2013). Replicative lifespan in budding yeast is defined as the number of daughter cells that a mother cell is able to produce and is therefore a measure of a cell’s replicative capacity at birth. By measuring lifespan in segregants from a cross between two genetically diverse strains of yeast, Kwan et al. were able to identify a point mutation at the rARS that (1) decreased origin activity at the rDNA; (2) increased origin activity elsewhere in the genome; and (3) prolonged replicative lifespan. They interpreted these results as reflecting competition for replication resources between the origins in two distinct compartments of the nucleus, namely the heterochromatic repetitive nucleolus or rDNA and the non-repetitive sequences that make up the bulk of the yeast genome. This competition is exacerbated by the perturbation of nucleolar heterochromatin caused by the absence of SIR2 (Yoshida et al. 2014; Foss et al. 2017). Most replication outside of the rDNA initiates from approximately 300 replication origins, and thus the 150 origins at the rDNA make up approximately one third of total origins. Increased replication resources directed toward the rDNA might therefore be expected to create a significant strain on replication in the rest of the genome.

Replication origins are known to compete for limiting resources at the level of origin firing. Broadly speaking, replication origin activity is regulated at two levels, namely origin licensing and origin firing (Remus and Diffley 2009). Origin licensing refers to loading of the replicative helicase in G1, and only origins that have been licensed are capable of being used in the subsequent S phase. Licensing occurs when an Origin Recognition Complex, Orc1-6, that is bound to an 11- 17 base pair so-called ARS consensus sequence (ACS) within the origin loads two hexameric helicases, each consisting of a ring of subunits, MCM2-7, that encircles the DNA (Bell and Stillman 1992; Francis et al. 2009; Bell and Botchan 2013; Deegan and Diffley 2016; Deegan et al. 2016; Douglas et al. 2018; Lewis et al. 2022). The mechanism by which Orc recognizes origins and loads an MCM double-hexamer has been elucidated in great detail (Bell and Labib 2016; Costa and Diffley 2022; Hu and Stillman 2023). DNA replication initiates from only a subset of these licensed origins, namely those where the helicase complex is activated, a process referred to as origin firing (Rhind and Gilbert 2013). Origin firing is triggered by activation of two kinase signaling pathways, DDK and CDK, which then cause each hexameric MCM helicase to encircle a single strand of DNA and track in opposite directions in front of the replisome, thereby establishing bidirectional replication (Heller et al. 2011; Larasati and Duncker 2016; Li et al. 2023). Firing factors are not present in sufficient quantities to activate all licensed origins simultaneously, and thus origins are activated in waves, with firing factors being recycled until the entire genome has been replicated (Mantiero et al. 2011; Tanaka et al. 2011). Determining the criteria by which origins are prioritized for activation remains an area of active research.

Importantly, if the competition between the origins at the rDNA and those in the rest of the genome is tilted in favor of the rDNA, either by a point mutation in the rARS or by deletion of SIR2, replication gaps between distantly spaced origins outside of the rDNA persist late into S phase, and genome duplication may remain incomplete when the cell enters mitosis (Foss et al. 2017). Indeed, the observation that weakening the rARS by a point mutation can suppress the short replicative lifespan characteristic of sir2 mutants suggests that these replication gaps may be the ultimate cause of "replicative death". Teleologically, one can rationalize the prioritization of replication of the unique areas of the genome over the repetitive and heterochromatic areas, because (1) there are mechanisms available for repair in repetitive regions, like single-strand annealing, that are not available in unique regions, and they may therefore be able to survive lesions created when mitosis precedes the completion of replication; and (2) the cell is able to tolerate variation in copy number at the rDNA, but not deletion of protein-coding sequences in unique regions. Consistent with this view of the utility of assigning a low priority to replication of repetitive sequence, heterochromatic regions are also late-replicating in mammalian cells, and disease states have been associated with disruption of this pattern (Donley and Thayer 2013).

Despite its biological importance, the mechanism by which Sir2 suppresses replication initiation is poorly understood. Clues to Sir2’s role in replication may be gleaned from its role in transcription, although here, too, our understanding is limited. Sir2’s ability to repress transcription at the silent mating type loci has been postulated to follow from its role in chromatin compaction. Experiments demonstrating SIR2-dependent protection of heterochromatin from digestion with endonucleases or methylation by bacterial DNA methyltransferases support the notion that DNA in heterochromatin is generally less accessible (Gottschling 1992; Singh and Klar 1992; Loo and Rine 1994; Ansari and Gartenberg 1999). Thus, a widely publicized model attributes Sir2’s ability to inhibit transcription to blocking access of RNA PolII to the DNA, although a simple model of indiscriminate steric occlusion cannot account for this phenomenon (Chen and Widom 2005; Lynch and Rusche 2009; Kitada et al. 2012; Bondra and Rine 2023). Whether Sir2’s effect on replication is similarly the result of chromatin compaction limiting access to enzymes involved in DNA metabolism is not known.

A study examining the distribution of MCM helicase complexes at the rDNA suggests that, rather than a global manifestation of lack of accessibility of the rDNA, the early activation of replication at the rDNA in the absence of Sir2 may reflect a more localized phenomenon (Foss et al. 2019). In this study, MCM binding was assessed at nucleotide-level resolution by fusing a subunit of MCM to micrococcal nuclease (MNase), a technique known as Chromatin Endogenous Cleavage, or "ChEC" (Figure 1B) (Schmid et al. 2004; Zentner et al. 2015). Permeabilization of cells with detergent followed by addition of calcium activates MNase, thereby generating approximately 65 base-pair fragments of DNA under MCM double-hexamers that can be identified by high- throughput sequencing. This analysis revealed that, while deletion of SIR2 did not elicit an obvious change in the overall levels of licensing at the rDNA, it led to a dramatic redistribution of MCM helicase complexes: In wild-type cells, the helicase complexes are located almost exclusively immediately adjacent to a high occupancy well-positioned nucleosome approximately 130 base pairs to the left of the ACS. In contrast, the MCM complexes in sir2 mutants are largely displaced to the right into an area with lower nucleosome occupancy (see Figure 1C and Figure 1-figure supplement 1). This displacement appears to be caused by the de-repression of C-pro transcription in the absence of SIR2, which results in RNA PolII pushing the MCM double-hexamer into the relatively nucleosome-free region. This observation is provocative in light of studies demonstrating that the arrangement of MCM-proximal nucleosomes can have a dramatic effect on MCM activation (Azmi et al. 2017). Specifically, by using chromatin remodeling enzymes (CREs) to create a variety of nucleosomal arrangements on an origin-containing fragment of DNA, Azmi et al. demonstrated that activation of an MCM double-hexamer loaded at that origin could be inhibited by nucleosomes arranged by specific CREs. In another study (Chacin et al. 2023), a mutant version of Orc1 that retains the ability to load MCM but alters phasing of MCM adjacent nucleosomes was shown to inhibit DNA replication in vivo and in vitro, consistent with a critical role for nucleosome organization at origins in the initiation of replication.

Taken together, the observations above led us to hypothesize that (1) relocation of MCM helicases from regions of high nucleosome occupancy to nearby regions of low nucleosome occupancy promotes early activation of replication origins in sir2; (2) this local difference in nucleosome occupancy is the result of modulation of nucleosomal architecture by one or more CREs; and therefore (3) screening CREs for the ability to suppress the early rDNA replication phenotype in sir2 mutants might allow us to identify the relevant CRE. In this report, we show that the displaced MCM complexes fire before their nondisplaced counterpart and that increased nucleosome occupancy near the displaced MCMs in the absence of Fun30 CRE suppresses their early activation. This is the first in vivo demonstration of the modulation of an origin activity by a specific chromatin remodeler in vivo, and has implications for the broader relationship between transcription and replication timing.

Results

Fun30 chromatin remodeling activity is required for early rDNA replication in sir2 mutant

To identify CREs required for the early replication of the rDNA in sir2 mutants, we needed a method to determine replication timing, and for this we turned to S-seq. In this widely used technique, G1 and S phase populations of cells from a log phase culture are sorted using flow cytometry and their DNA subjected to genomic sequencing (Muller et al. 2014; Batrakou et al. 2020). Replication timing can then be inferred from read depths in the S phase fraction, because regions that replicate early in S phase will be present in two copies in most cells, while regions that replicate late will be present mostly as a single copy. S phase read depths are normalized to those in G1 to account for differences in DNA copy numbers, which is particularly important for repetitive loci whose copy number can vary, such as rDNA. As shown in Figure 2A, which depicts the S-seq profile of chromosome 12, relative replication timing (Trel) values using this approach typically range from 1.4 to 0.8 for early- and late-replicating regions, respectively, with a value of 1 indicating that a locus replicates at the same time as the genome-wide average (see Materials and Methods for specific calculation). The late replication of the rDNA, defined as coordinates 450,000 - 470,000 on chrXII, in wild type is reflected in a Trel value of 0.89, whereas this ratio increases to 1.2 upon deletion of SIR2 (Figure 2A,B). Using this method, we showed that a point mutation in the ACS of the rARS suppresses early replication of the rDNA in sir2, while promoting replication in the rest of the genome (Kwan et al. 2013; Foss et al. 2017). To identify potential chromatin remodelers required for early rDNA replication in sir2, we measured rDNA replication timing in double-mutants in sir2 and each of 10 non-essential CRE subunits (Chd1, Fun30, Htl1, Irc5, Isw1, Isw2, Nhp10, Swp82). Deletion of FUN30, but no other gene, delayed replication of the rDNA in a sir2 background, shifting the Trel value from 1.2 to 1.06, while having had no significant effect in SIR2 cells (0.89 for WT and 0.91 for fun30) (Figure 2B).

S-Seq applied to rDNA replication timing and copy number determination.

A. S-seq replication profile of Chromosome XII. The region in the middle noted as the rDNA has been collapsed. Note that the rDNA replicates much earlier in sir2 (16559) (orange) than in wild type (16535) (blue). B. rDNA replication timing in double mutants between sir2 and various chromatin remodeling enzymes: nhp10 (16729), chd1 (16725), swp82 (17061), isw1 (16724), htl1 (17059), swr1 (16728), isw2 (16673) and irc5 (16723). fun30-S20A S28A (17113 and 17114) is non-phosphorylatable, and fun30-K603R (17345 and 17346) is catalytically inactive. The S to G1 ratio values (mean±SD) were 0.90±.0.02 for WT, 1.04±0.04 for fun30sir2 (16909), 1.01±0.02 for fun30K603R and 1.21±0.03 for sir2 (p<0.001 for WT vs sir2, sir2 vs fun30 sir2 and sir2 vs fun30K603Rsir2 using t-test). C. Effect of sir2 and fun30 mutation on rDNA size, as determined from the fraction of G1 sequencing reads that arise from 450-470 kb on chrXII. The values (% mean±SD) were 9.4±1.6 for WT, 4.6±1.2 for fun30sir2, 4.8±0.3 for fun30K603R and 9.2±1.8 for sir2 (p<0.001 for WT vs sir2, sir2 vs fun30sir2 and sir2 vs fun30K603R sir2 using t-test)

Because Fun30 can form a complex with the replication scaffold protein Dpb11, it was possible that Fun30 played a structural rather than an enzymatic role in regulating rDNA replication. However, we found that a point mutation that abolishes Fun30 ATPase activity, K603R (Neves-Costa et al. 2009; Eapen et al. 2012), also suppresses early rDNA replication (Figure 2B), arguing that Fun30 instead promotes rDNA replication through its enzymatic activity in nucleosome remodeling. Fun30 is also known to play a role in the DNA damage response; specifically, phosphorylation of Fun30 on S20 and S28 by CDK1 targets Fun30 to sites of DNA breaks, where it promotes DNA resection (Chen et al. 2016; Bantele et al. 2017). To determine whether the replication phenotype that we observed might be a consequence of Fun30’s role in the DNA damage response, we tested non-phosphorylatable mutants for the ability to suppress early replication of the rDNA in sir2; these mutations had no effect on the replication phenotype (Figure 2B), arguing against a primary role for Fun30 in DNA damage repair that somehow manifests itself in replication.

Fun30 chromatin remodeling activity is required for maintaining rDNA array size in sir2 mutant

We next wanted to determine whether FUN30 affects rDNA array size. The repetitive nature of the rDNA allows it to fluctuate in size, and mutations that impair DNA replication, either globally or in a manner specific to the rDNA, are associated with shrinkage of the rDNA. For example, the rDNA in a wild-type strain in the S288c background consists of 150 copies, but a point mutation at the ARS consensus sequence (ACS) that impairs Orc binding causes the rDNA to shrink to just 80 copies (Kwan et al. 2013). rDNA shrinkage in response to global replication defects is thought to reflect the fact that the rDNA, which typically comprises approximately 10% of genomic DNA, imposes a significant replicative burden on the cell that can be ameliorated by decreasing the number of repeats (Salim et al. 2017).

We determined rDNA size by assessing the fraction of DNA sequencing reads from a sonicated G1 population that arose from the rDNA. While deletion of FUN30 had no notable effect on rDNA size in a SIR2 strain, it caused a 50% reduction in a sir2 mutant, from 9.2% to 4.2% (Figure 2C). rDNA size reduction of a similar magnitude was observed in the ATPase-deficient Fun30K603R mutant, but not in the non-phosphorylatable Fun30S20AS28A mutant, mirroring the disparate effect of the two alleles on rDNA replication timing. The reduction of the rDNA size in sir2 fun30 mutants was confirmed by qPCR analysis of DNA isolated from log cultures, and by Southern blots (Figure 2-figure supplement 1 and Figure 3). Based on measurements in log-growing cells, we estimated the rDNA size in fun30 sir2 mutants to be ∼40 repeats (Figure 2-figure supplement 1). This places it in a category comparable to the shortest rDNA arrays in a collection of DNA replication mutants noted for their contraction of the rDNA (Salim et al. 2017), thus emphasizing the severity of the replicative disadvantage conferred by the deletion of FUN30 when Sir2 function is absent.

Licensing at the rARS, as determined by Southern blot.

Activation of MCM-MNase in G1-arrested WT (16747), sir2 (16769), sir2fun30 (17257) and fun30 (17256) cells at the rARS with calcium will eliminate the 3.5 kb XmnI fragment (upper panel). PIK1, a single-copy gene in which we detect no MCM binding, is used as a loading control. Normalized ARS1200 band intensity at 15m is expressed relative to time 0. Quantitation of the uncut band was used to infer relative rDNA array size in sir2 fun30 mutant at 0.35 relative to WT.

Alterations of rDNA size are promoted by the replication fork barrier (RFB), which imposes a unidirectional replication fork arrest to prevent head-to-head collisions between DNA polymerase and RNA polymerase I while transcribing 35S units. Replication forks that are stalled at the RFB are prone to breakage, which can lead to unequal sister chromatid exchange and fluctuations in rDNA array size. These rDNA size fluctuations are exacerbated in sir2 mutants, which exhibit an increased rate of recombination, resulting in an unstable rDNA array size. Deletion of FOB1, which abolishes RFB function, stabilizes rDNA size by reducing,though not eliminating,the frequency of recombination-initiating double-strand breaks.

We have interpreted the shrinkage of the rDNA caused by deletion of SIR2 and FUN30 to be a specific example of the general phenomenon demonstrated by Salim et al. (Salim et al. 2017), i.e. that mutations that impair replication generate selective pressure that favors rDNA shrinkage. According to this interpretation, while deletion of FOB1 is expected to retard the transition to reduced array sizes, it should not prevent it, and thus rDNA shrinkage should not be FOB1- dependent.

To determine whether the reduction of rDNA size in sir2 fun30 mutants was dependent on RFB activity and FOB1, we deleted FUN30 in a sir2 fob1 strain with 160 repeats. Using this strategy, we obtained fob1 sir2 fun30 transformants with over 150 rDNA repeats.. If rDNA shrinkage does not require Fob1, it should occur, albeit slowly, in fob1 mutants. To assess this, we continuously grew several independent fob1 sir2 fun30 colonies, along with control fob1 sir2, and fob1 fun30 strains for 120 generations. We found that, unlike the control fob1 strains, which maintained a stable rDNA size, all four fob1 sir2 fun30 cultures reduced their rDNA arrays size to fewer than 100 copies by 120 division (Figure 2-figure supplement 2) . These results demonstrate that, although FOB1 affects the kinetics of rDNA size reduction in sir2 fun30 strains, the reduced rDNA array size upon FUN30 deletion size does not depend on FOB1. Thus, suppressed rDNA replicaton drives the selection for reduced rDNA array size in both RFB-proficient and RFB- deficient fun30 sir2 cells. Consistent with this notion, our S-seq results showed delayed rDNA replication timing in fob1 sir2 fun30 cells with 150 rDNA repeats, compared to their fob1 sir2 counterparts (Figure 2-figure supplement 3). In summary, we conclude that Fun30 nucleosome remodeling activity promotes rDNA replication and is required for maintaining rDNA size in sir2 mutants. We next aimed to determine whether Fun30 exerts this effect by reducing MCM loading or the ability of loaded MCMs to fire.

High level of origin licensing at rDNA origins in WT, sir2 and sir2 fun30 mutants

It has been estimated that ∼20% of potential rDNA origins are activated during any single S-phase in wild-type cells (Brewer and Fangman 1988; Linskens and Huberman 1988). Furthermore, single-molecule studies suggested that the fraction of active origins is increased in sir2 mutant cells (Pasero et al. 2002; Yoshida et al. 2014). To determine whether SIR2 and FUN30 modulate rDNA replication by changing origin licensing, i.e. the fraction of origins that are bound by MCM helicase complexes prior to the beginning of S phase, we assessed licensing in WT, sir2, fun30, and sir2 fun30 cells by measuring DNA cleavage by MCM2-MNase using Southern blots (Foss et al. 2021): Digestion with XmnI yields a 3.5 kb fragment that spans the rARS (Figure 3), and, if the rARS is licensed, activation of MNase by addition of calcium will eliminate this fragment, thus making it undetectable when using a probe upstream of the rARS. As a control, we used a fragment that, based on our MCM-ChEC data, should not contain MCM complexes. Activation of MNase in all 4 genotypes led to an ∼60-85% decrease in the XmnI fragment (Figure 3 and Figure 3 figure supplement 1), therefore we conclude that (1) licensing is high, at ∼60-85%; and (2) a much higher fraction of origins are licensed than are activated, making it unlikely that significant regulation of origin activity occurs at the level of licensing. This is the first quantitative assessment of licensing at the rDNA in either wild type or sir2.

Fun30 promotes firing of rDNA origins and suppresses firing of non-rDNA origins in sir2 mutants

Our observation that FUN30 status had no obvious effect on rDNA origin licensing suggested that Fun30 might instead advance rDNA replication in sir2 mutants by promoting rDNA origin firing. To evaluate this possibility, we measured the effect of FUN30 status on origin activity at the rDNA using 2D gels. We grew cells to log phase, arrested them for 1.5 hours in alpha factor and collected them at different time points after release into 200 mM hydroxyurea. Prior to separation on 2D gels, DNA was digested with NheI, which releases a 4.7 kb rARS-containing linear DNA fragment at the internal rDNA repeats (1N) and a much larger, 24.4 kb single-rARS-containing fragment originating from the rightmost repeat (Figure 4A). In 2D gels, active origins generate replication bubble arc signals, whereas passive replication of an origin appears as a y-arc (Figures 4B). Having a signal emanating from a single-rARS-containing fragment (Figures 4A and 4B) simplifies the comparison of rDNA origin activity in strains with different numbers of rDNA repeats, such as in sir2 vs sir2 fun30 mutants. Origin activity is expressed as a ratio of the bubble to the single- ARS signal, effectively measuring the number of active rDNA origins per cell at a given time point.

Deletion of FUN30 suppresses origin activity at rDNA and promotes it elsewhere in the genome.

A. Schematic of the right end of the rDNA locus depicting Nhe1 cut sites. Digestion with Nhe1 releases multiple copies of a 4.7 kb rARS-containing internal fragment (1N) and a single copy of a 24.4 kb fragment that contains the right-most rDNA origin. B. Diagram of replication structures detected by 2D gel. C. Replicas of 2D gels showing replication at the rDNA locus 30 minutes after release from alpha factor into HU in sir2 fob1 (17564) or sir2 fun30 fob1 (17556) strains with short (29 copies and 35 copies, respectively) rDNA arrays. Numbers indicate the ratio of bubble arc to single copy ARS signal, reflecting the number of activate origins per cell. The average number of activate origins (mean±SD) for sir2 fob1 is 16.9±2.2 vs 4.2±0.7 (P<0.05 by t-test). D. 2D gels showing replication at the rDNA. Cells were arrested in G1 with alpha factor and then released into medium containing 200 mM HU. Quantitation of the ratio of bubble arc to 24.4 kb single copy rARS is shown below. * indicates significant differences (p ≤ 0.05 as determined by comparing combined signals from all 4 time points by Student’s t-test). Strain numbers used were as follows: wild type (16747), sir2 (16769), sir2 fun30 (17257) and fun30 (17256). E. Replication of ARS305 was examined as in D, but the 1N spot was used for normalization and ratios of bubble arc to 1N spot were normalized to this ratio for the 20 minute time point in wild type. F. EdU incorporation at 111 early origins 1 hour after release from G1 into medium containing 200 mM HU. Total genome-wide read counts in each sample were normalized to the number of reads in the sample with the highest total, thereby normalizing numbers to genome-wide incorporation of EdU. Each dot represents a single origin, with read depths summed across a 5 kb window centered on the MCM binding site within each origin. Points are plotted according to EdU signal in wild type on the x axis and in the mutant in question on the y axis. Suppression of origin activity is reflected in points dropping below the dotted line at 45°. The rARS is circled in green. Strain numbers used were as follows: wild type (17265), sir2 (17271), sir2 fun30 (17279) and fun30 (17281).

As seen previously, deletion of SIR2 increased the number of activated rDNA origins, while deletion of FUN30 suppressed this effect (Figure 4D). When analyzed in aggregate at 20, 30, 60 and 90 minutes following release into HU, the average number of activate rDNA origins in sir2 mutant was increased 6.3-fold compared to those in WT (5.0±2.3 in sir2 vs 0.8±0.4 in wt, p<0.05 by two-tailed t-test), and this increased number was reduced upon FUN30 deletion (1.3±0.7 in sir2 fun30, p<0.05 by 2 tailed t-test vs sir2, NS for comparison to WT). Deletion of FUN30 in a strain with a wild-type copy of SIR2 did not have asignificant effect. We conclude that deletion of FUN30 partially reverses the accelerated activation of ribosomal origins caused by deletion of SIR2.

The overall impact of FUN30 deletion on rDNA origin activity in a sir2 background is expected to be a composite of two opposing effects: a suppression of rDNA origin activation and increased rDNA origin activation due to reduced rDNA size. To evaluate the effect of FUN30 on rDNA origin activation independently of rDNA size, we generated an isogenic set of strains in a fob1 background, all of which contain ∼30 copies of the rDNA repeats. (Deletion of FOB1 is necessary to stabilize rDNA copy number.) Comparing rDNA origin activity in sir2 versus sir2 fun30 genotypes, we observed a robust and reproducible reduction in rDNA origin activity upon FUN30 deletion (Figure 4C). This finding confirms that the FUN30 suppresses rDNA origin firing in sir2 backgrpound independently of both rDNA size and FOB1 status.

We and others have previously shown that deletion of SIR2 not only advances replication at the rDNA but also retards replication elsewhere in the genome (Kwan et al. 2013; Yoshida et al. 2014). The absence of SIR2 has been linked to decreased activation of less robust early replication origins, likely due to their heightened sensitivity to overall levels of replication-initiating factors. To ascertain the specificity of the interaction between FUN30 and SIR2, we investigated whether deleting FUN30 also mitigates the delayed activation of weak early origins in a sir2 mutant background. We focused on the activation of ARS305, an origin we previously identified as delayed in a sir2 mutant using 2D gel analysis. Overall, the comparison of origin activity at ARS305, measured as bubble arc to 1N signal, with that at rDNA demonstrates a reciprocial relationship between activity at ARS305 and rDNA across the four genotypes (Figure 4E). As seen previously, deletion of SIR2 suppresses ARS305, whereas additional deletion of FUN30 increased activity close to that of wild type. Deletion of FUN30 in a SIR2 background also showed a trend of increased origin activation, though this difference was not statistically significant.

Delayed initiation of genomic origins following SIR2 deletion has been previously demonstrated on a global scale as a decline in the incorporation of a labeled uridine analogue, bromodeoxyuridine, at weak early origins in sir2 cells compared to wild-type cells upon their release from G1 into HU (Yoshida et al. 2014). We employed a similar approach to investigate whether the deletion of FUN30 mitigates this effect, utilizing a different labeled nucleotide, ethynyl-deoxy-uridine (EdU), which can be isolated through click chemistry rather than immunoprecipitation. Wild-type (WT), sir2, sir2 fun30, and fun30 cells engineered for the uptake and phosphorylation of the analogue (Viggiani and Aparicio 2006) were arrested in G1 using alpha factor and collected 60 minutes after release into medium containing HU and EdU. DNA extracted from these cells was sonicated and subjected to click chemistry with biotin azide to covalently link biotin to EdU incorporated into DNA. Biotinylated DNA was then purified and sequenced. The impact of SIR2 and FUN30 status on replication timing was assessed by graphing the EdU signal at 111 early origins in WT cells on the x-axis against the signal in sir2, fun30, and sir2 fun30 cells on the y-axis (Figure 4F). Similar to observations with BrdU, we noted reduced EdU incorporation at a subset of early origins in sir2 mutants, coupled with increased EdU incorporation at rDNA origins. Deletion of FUN30 in a sir2 background partially restored EdU incorporation at early origins, concomitant with reduced EdU incorporation at rDNA origins. In particular, the median value of log10 of read depths at 111 early origins, as the data are shown in Figure 4F, dropped from 6.5 for wild type to 6.2 for sir2 but then returned almost to wild type levels (6.4) in sir2 fun30. The p value obtained by Student’s t test, comparing the drop in 111 origins from wild type to sir2 with that from wild type to sir2 fun30 was highly significant (<< 10-16) In contrast, FUN30 deletion in the WT background did not reduce EdU incorporation at genomic origins (median 6.6). These findings highlight that FUN30 deletion-induced suppression of rDNA origins in sir2 is accompanied by the activation of genomic origins.

Early firing of displaced helicase complexes

While our results above demonstrate that deletion of SIR2 promotes firing of rDNA origins and that additional deletion of FUN30 suppresses this phenotype, they lack the resolution to distinguish replication arising from the displaced versus the nondisplaced MCM complexes. Other methods for measuring origin activation, such as generation of single-stranded DNA or incorporation of nucleoside analogs, also suffer from this limitation. Distinguishing between firing of the displaced versus the nondisplaced MCM complexes is important to test our hypothesis that displacement of the helicase complex into a relatively nucleosome-free region is responsible for its early activation. We therefore sought to devise a method that can differentiate activation of MCM complexes that are only ∼150 bp apart, i.e. the approximate distance between displaced and nondisplaced MCMs. We reasoned that, because MCM-ChEC has the resolution to distinguish closely spaced helicase complexes, the differential disappearance of such closely spaced footprints might permit us to distinguish their relative activation times.

To explore the possibility that the disappearance of the MCM-ChEC signal could be used to distinguish the firing of the displaced versus nondisplaced populations of MCM, we first asked whether the pattern of disappearance of the MCM signal recapitulates that of origin firing. In particular, while late origins are inhibited by hydroxyurea (HU), early origins are not; therefore, we asked whether this phenomenon was reflected in MCM-ChEC. To do this, we arrested cells in G1 using alpha factor and then released them into medium containing hydroxyurea (HU), taking samples for ChEC analysis at 15, 30, 45, 60 and 90 minutes (Figure 5A). HU inhibits ribonucleotide reductase and thereby dramatically slows S phase, making it easier to follow replication kinetics, but the relative order of firing of different origins remains unchanged. In Figure 5B, we plot the MCM-ChEC signal at 111 early origins and 101 late origins, with G1 on the x-axis and S phase on the y-axis. While we observed little difference in the signal abundance between early and late origins just 15 minutes after release from alpha factor, before origins have had a chance to fire, this difference increases over the time course, with early origins dropping more than late. MCM signal quantification demonstrates that the differences in MCM abundance between early and late origins increase over the time course in HU (Figure 5-figure supplement 1). Using chromatin immunoprecipitation followed by sequencing (ChIP-seq) for FLAG-tagged MCM2 in cells arrested in alpha factor and those released into HU for 60 minutes, we observed a similar reduction of the MCM signal at early compared to later origins, confirming our MCM- ChEC results (Figure 5-figure supplement 2). We conclude that the disappearance of the MCM- ChEC signal as cells progress through S phase can be used to assess origin firing.

Disappearance of MCM2-ChEC signal can be used to monitor origin firing.

Earlier version: A. Cells were synchronized in G1 phase using alpha factor for 1.5 hours before being released into media supplemented with hydroxyurea (HU). Cells were harvested at various time points post-release (15-90 minutes) and subjected to MCM2-ChEC analysis. B. Decrease in MCM2-ChEC signal in wild type (16747) for 111 early (orange) and 101 late (blue) origins. MCM2-ChEC signal was quantified over 200 base pair windows centered on the MCM binding site within each origin. Each point is plotted according to MCM2-ChEC signal at the time point in question on the y axis and the corresponding signal in G1 on the x axis, thus decrease in signal appears as a drop below the 45° diagonal. C. MCM signal at the rDNA. MCM2-ChEC signal in G1 appears predominantly at the location indicated by the blue rectangle seen in wild type (16747) and fun30 (17256), whereas it is spread across both the blue and orange rectangles in the absence of Sir2. The displaced MCM2-ChEC signal in sir2 (16769) disappears more rapidly than its non-displaced counterpart, and this effect is suppressed by fun30 (17257). Decrease of MCM2-ChEC signal at the rARS is accompanied by increased signal to the right.

Revised version: A. Cells were synchronized in G1 phase using alpha factor for 1.5 hours before being released into media supplemented with hydroxyurea (HU). Cells were harvested at various time points post-release (15-90 minutes) and subjected to MCM2-ChEC analysis. B. Decrease in MCM2-ChEC signal in wild type (16747) for 111 early (orange) and 101 late (blue) origins. MCM2-ChEC signal was quantified over 200 base pair windows centered on the MCM binding site within each origin. Each point is plotted according to MCM2-ChEC signal at the time point in question on the y axis and the corresponding signal in G1 on the x axis, thus decrease in signal appears as a drop below the 45° diagonal. C. MCM signal at the rDNA. MCM2-ChEC signal in G1 appears predominantly at the location indicated by the blue rectangle seen in wild type (16747) and fun30 (17256), whereas it is spread across both the blue and orange rectangles in the absence of Sir2. The displaced MCM2-ChEC signal (orange rectangle) in sir2 (16769) disappears more rapidly than its non-displaced counterpart, and this effect is suppressed by fun30 (17257). Decrease of MCM2-ChEC signal at the rARS is accompanied by increased signal to the right.

We then turned our attention to the MCM signals at the rDNA origin; here the presence of the replication fork block approximately 1.5 kb from the sites of MCM binding provides a fortuitous means of confirming our interpretation, since the loss of the MCM-MNase signal at the origin should appear as a reciprocal increase where the replisome stalls. We first used MCM2-Flag ChIP in WT and sir2 cells after release into HU to confirm this prediction. In both WT and sir2 cells, the MCM can be found at the rDNA origin in alpha factor-arrested cells as expected (Figure 5 figure supplement 3). However, at 60 minutes following release into HU, unlike in WT cells where the MCM signal remains at its loading site, in sir2 mutants, we can see that a large fraction of the signal has now redistributed from its loading site and accumulated left of the RFB, consistent with early activation of rDNA origins in sir2 cells and the arrest of the replicaiton forks at the RFB. While these results demonstrate that ChIP can be used to monitor the repositioning of the MCM signal at a kb scale and origin activation, its resolution is inadequate for small-scale changes in MCM signal distribution, which is needed to monitor the proposed differences in activation of repositioned and non-repositioned MCM complexes. We therefore turned to ChEC to address this question.

The MCM-ChEC signal in wild-type cells remained at their initial loading throughout the duration of the HU experiment (blue box in Figure 5C) as we have seen by ChIP in WT cells at 60 minutes. (Note that 200 mM HU causes an approximately 20-fold slowing of S phase, therfore a 60 minute time point is equivalent to <5 minutes without HU, i.e. prior to the time when the rDNA origin fires in wild-type cells). Like in wild type, this nondisplaced population of MCM in sir2 remains relatively constant (blue box); however, in striking contrast, the displaced population of MCM in sir2 (orange box) shows a continual decrease over the course of the experiment, and there is a concomitant increase in signal in the region upstream of RFB that is absent in wild type. Furthermore, both the decrease in the displaced signal (orange box) and its corresponding increase upstream of the RFB were suppressed by deletion of FUN30. The changes in the relative abundance of displaced and nondisplaced MCMs during the time course, as well as the repositioning of the MCM signal to the region upstream of RFB across different genotypes, were highly reproducible (Figure 5-figure supplement 4).

Taken together, these observations provide strong support for the hypothesis that the displaced population of MCM in sir2 cells fires more readily than its nondisplaced counterpart and that the increased firing of the displaced MCMs is suppressed by FUN30 deletion.

It is notable that, although a small fraction of the total G1 signal in wild type arises from the displaced location, this minor signal likewise disappears more rapidly than its nondisplaced counterpart (Figure 5-figure supplement 5). This suggests that some feature of the immediate chromatin environment may be responsible for the difference in firing times at the two locations, and that this difference is present in both wild type and sir2; in other words, the fact that the displaced helicase complex is activated before the nondisplaced complex regardless of SIR2 status is consistent with a model in which sir2 mutants replicate their rDNA early by virtue of displacing much of their loaded helicases to a region where they are more apt to fire, rather than by altering the chromatin environment more globally to affect all replicative helicases in the nucleolus.

Earlier version: Finally, although deletion of FUN30 could suppress replication initiation at the rDNA either by inhibiting the firing of the active (displaced) MCM complex or by preventing its displacement to the "activating location" in the first place, our results make it clear that suppression occurs via the former mechanism rather than the latter. Indeed, deletion of FUN30 actually skews the distribution to favor the displaced location over the nondisplaced one, and this skewing can be observed in both sir2 and SIR2. As we have previously shown, rightward displacement of MCMs is accomplished by the derepression of C-pro transcription and the consequent pushing by RNA polII, suggesting that deletion of FUN30 does not act via repression of C-pro transcription. Consistent with both this expectation and a previous report demonstrating that fun30 mutants are deficient in transcriptional silencing (Neves-Costa et al. 2009), C-pro RNA levels were approximately twice as high in sir2 fun30 cells compared to those in sir2 cells when adjusted for differences in rDNA size (Figure 5-figure supplement 6). A combination of increased MCM repositioning and reduction of MCM signal at the RFB observed in sir2 fun30 mutants also confirms that the accumulation at the RFB is a reflection of replication fork stalling rather than RNA PolII-mediated repositioning of the MCMs. Taken together, our findings indicate that FUN30 deletion specifically hampers the activity of repositioned MCMs when silencing mechanisms are compromised.

Revised version: Finally, although deletion of FUN30 could suppress replication initiation at the rDNA either by inhibiting the firing of the active, repositioned MCM complex or by preventing MCM repositioning to the "active location" in the first place, our results suggest that suppression occurs through the former mechanism. Consistent with previous reports that fun30 mutants are deficient in transcriptional silencing (Neves-Costa et al. 2009), C-pro RNA levels were approximately twice as high in sir2 fun30 cells compared to sir2 cells when adjusted for rDNA size (Figure 5—figure supplement 6).

Added in the revised version: Moreover, deletion of FUN30 shifts the distribution toward the repositioned MCM location over the non-repositioned one in G1 cells (Figure 5C), aligning with the increased C-pro transcription observed in fun30 mutants. This shift is evident in both sir2 and SIR2 cells. Despite the increased transcription-mediated repositioning in sir2 fun30 cells compared to sir2 cells during G1, repositioned MCM persists longer in sir2 fun30 cells than in sir2 cells after release into HU. Additionally, sir2 fun30 mutants exhibit reduced MCM accumulation at the RFB compared to sir2 mutants after release into HU, supporting the conclusion that MCM disappearance in HU reflects origin activation rather than transcription-mediated displacement.

Fun30 maintains low nucleosome occupancy near repositioned MCMs

Our findings demonstrate that MCMs, when situated at their loading site adjacent to a high- occupancy +1 nucleosome, exhibit reduced firing propensity. In contrast, when these MCMs are displaced to a neighboring region with lower nucleosome occupancy, their firing propensity significantly increases. This heightened firing tendency is dependent on Fun30 CRE and its remodeling activity. This suggests that the absence of FUN30 could elevate nucleosome occupancy around the relocated MCMs, consequently inhibiting their ability to fire.

To evaluate this possibility, we examined the effect of FUN30 deletion on nucleosome occupancy at rDNA origins in WT and sir2 background using MNase-seq in G1 arrested cells. Chromatin treatment with MNase digests DNA that is not bound by proteins, such as nucleosomes and transcription factors. Nucleosomes, which are wrapped by approximately 150 base pairs of DNA, tend to protect fragments in the 150-200 base pair range, and thus profiling fragments in this size range is widely used to assess nucleosome distribution. Plotting nucleosome midpoints at the rARS in G1 cells (Figure 6A and Figure 6-figure supplement 1) reveals that nucleosome occupancy at the +2 (green box) and +3 (red box) positions is higher in the sir2 fun30 double mutant than in the sir2 single mutant. Furthermore, this effect of Fun30 is not limited to the sir2 background, since nucleosome occupancy is also modestly increased in fun30 cells compared to wild type cells. These observations support the notion that it is the relatively nucleosome-free nature of the region into which MCMs are displaced in sir2 that make the displaced complexes more prone to activation. Furthermore, this suggests that it could be particularly informative to assess nucleosome occupancy specifically adjacent to loaded MCM complexes, as described below.

Nucleosome occupancy assessed by MNase-seq and MCM2-ChEC.

A. Analysis of nucleosome occupancy at rDNA origins using MNase-seq revealed a consistent high occupancy of the +1 nucleosome across WT (16747), sir2 (16769), sir2 fun30 (17257) and fun30 (17256), which served as our normalization reference. Occupancy at the +2 (green box) and +3 (red box) positions was increased by deletion of FUN30 in both sir2 and SIR2. MCM2-ChEC signal was quantified specifically from the 151-200 base pair (nucleosome) size range. B. Analysis of nucleosome occupancy with MCM2-ChEC (see Figure 1B) reveals nucleosome occupancy in that subset of cells and rDNA repeats in which MCM is present. Deletion of FUN30 leads to increased occupancy at the +2 and +3 positions in a sir2 background. MCM2-ChEC signal was quantified specifically from the 151-200 base pair (nucleosome) size range.

It has previously been observed, both with MCM and with various transcription factors, that the ChEC technique can reveal not only the binding site of the protein to which the MNase is fused, but also, if present, that of the adjacent nucleosome (Zentner et al. 2015; Foss et al. 2019). This is due to the fact that the entrance and exit sites for DNA wrapped around a nucleosome are in close physical proximity, despite the fact that they are separated by approximately 150 base pairs, and thus the MNase fusion protein adjacent to a nucleosome often cleaves both sites. MCM-ChEC, therefore, provides a powerful tool for measuring nucleosome occupancy immediately adjacent to the MCM helicase complex exclusively in those rDNA repeats where those complexes have been loaded. By plotting the midpoints of 150-200 bp fragments released by MCM2-Mnase, we were able to determine that nucleosome occupancy at both the +2 and +3 nucleosomes, i.e. those adjacent to the displaced MCM complexes, are elevated in sir2 fun30 as compared to sir2 (Figure 6B and Figure 6-figure supplement 1). This finding indicates that Fun30 is required for maintaining the low nucleosome occupancy in the region into which the preponderance of MCM complexes have been displaced in sir2 mutant cells.

Discussion

In this study, we have attempted to shed light on the mechanism by which Sir2 represses replication origins at the rDNA. Sir2’s activity has been widely ascribed to the intuitive, but mechanistically imprecise, notion that "closed chromatin" is less accessible to both transcription and replication factors. Consistent with this model, deletion of SIR2 abolishes heterochromatic compaction and simultaneously increases the ability of methylases and restriction endonucleases to act on those sequences (Gottschling 1992; Singh and Klar 1992; Loo and Rine 1994; Weiss and Simpson 1998; Ansari and Gartenberg 1999; Ravindra et al. 1999). However, at least in its simplest form, non- specific steric exclusion of DNA metabolic enzymes cannot fully account for the suppression of transcription and replication at the rDNA. Steric arguments such as these are even less compelling when made for rDNA than for the silent mating type loci and telomeres, because chromatin compaction has been studied mostly in the context of the complete Sir complex (Sir2-4) (Gartenberg and Smith 2016). In contrast, Sir3, and Sir4 are not present at the rDNA (Straight et al. 1999). Moreover, rDNA is the most highly transcribed region of the genome, with transcription by PolI and PolIII accounting for approximately 80% of cellular RNA (Warner 1999; Woolford and Baserga 2013). These two polymerase complexes are similar in size to the PolII complex, thus size alone is not sufficient to selectively occlude PolII.

Additionally, deletion of SIR2 increases accessibility of restriction enzymes to heterochromatin at the HML locus by 1.5-fold, yet it increases PolII transcription by over 3000-fold (Holland 2002). This indicates that the magnitude of Sir2’s steric exclusion is insufficient to explain its effect on PolII transcription. Similar measurements in eukaryotic cells found little difference in accessibility of restriction enzymes and micrococcal nuclease between euchromatin and heterochromatin (Chereji et al. 2019).

In the current report, we argue that Sir2’s role in suppressing replication at the rDNA reflects local differences in the distribution of the MCM replicative helicase complexes rather than more global occlusion of these complexes from the rDNA. In particular, our results suggest that the key difference between wild type and sir2, with regard to replication initiation, is the location of the MCM complex, with the preponderance of that complex in wild type abutting a high-occupancy nucleosome while, in sir2, it is mostly in a relatively nucleosome-free region. Furthermore, we show that, even though this nucleosome-free region is less than 200 base pairs from MCM’s original location, the displaced helicase complex is significantly more prone to activation than its nondisplaced counterpart. Moreover, the chromatin-remodeling enzyme Fun30 is required to maintain the nucleosome-free character around the displaced MCM, and its deletion both increases nucleosome occupancy and decreases the ability of the displaced MCM to fire. Finally, as discussed below, we suggest that the increased proclivity of the displaced MCM to fire is not an anomaly specific to the physiology of sir2 cells, but is an integral feature of replication of the rDNA during normal S phase.

Our model for the manner in which Sir2 regulates rDNA replication depends critically on the notion that the local nucleosome landscape can influence the activation of MCM helicase complexes (Figure 7), and in vitro experiments suggest that this is, indeed, the case (Azmi et al. 2017): Azmi et al. monitored both licensing and firing in vitro, using a 3.8 kb nucleosomal template containing the ARS1 origin. They found that, while licensing was equivalent whether the template had been populated with nucleosomes using Swi/Snf or Rsc remodelers, on the one hand, or Isw1a, on the other, the loaded MCM complexes in the latter templates were more prone to activation. Furthermore, MCM complexes on Swi/Snf- or Rsc-remodeled templates could be made more prone to activation if subsequently treated with Isw1a, although the converse was not true. The salient difference between the two types of nucleosomal landscapes, however, was not obvious, as the distance between the MCM complex and the nearest nucleosome was not notably affected, although it is possible that the composition (e.g. H2A/H2B content) differs. It was also unclear how the nucleosome landscape affected activation of MCM complexes, though the relevant step in activation was determined to be the formation of the CMG complex with Cdc45 and GINS rather than the prior phosphorylation of the complex by DDK. In summary, while it is clear that the local nucleosome landscape can affect origin activity (Eaton et al. 2010; Belsky et al. 2015), the underlying mechanism remains murky.

Model for relationship between MCM location and replication timing at the rDNA.

Earlier version: MCM helicase complex (purple ovals) in wild type abuts the +1 nucleosome (blue cylinder) in G1, making it relatively resistant to activation. Deletion of SIR2 derepresses C-pro transcription (arrow pointing to the right), and RNA PolII pushes MCM complex to a nucleosome free area, where it is more prone to activation. Deletion of FUN30 in a sir2 mutant leads to increased nucleosome occupancy at the +2 position, adjacent to MCM complex, making this complex resistant to activation. Short red stretch of DNA (e.g. between +2 and +3 nucleosomes in top row) indicates ACS. Created with BioRender.com.

Revised version: MCM helicase complex (purple ovals) in wild type abuts the +1 nucleosome (blue cylinder) in G1, making it relatively resistant to activation. Deletion of SIR2 derepresses C-pro transcription (arrow pointing to the right), and RNA PolII pushes MCM complex to a nucleosome free area, where it is more prone to activation. Deletion of FUN30 in a sir2 mutant leads to increased nucleosome occupancy at the +2 position, adjacent to MCM complex, making this complex resistant to activation. Short red stretch of DNA (e.g. between +2 and +3 nucleosomes in top row) indicates ACS. "Ghost" Mcm complex in wild type indicates the small fraction of Mcm that is displaced even in wild type cells (see Figure 5-figure supplement 5). Created with BioRender.com.

Regardless of the reason for the increased proclivity of the displaced MCM to fire, the observation has implications for replication at the rDNA not only in sir2 mutants but also in wild type. We have previously shown that the rightward displacement of MCM double hexamers in sir2 is due to derepression of C-pro transcription (Foss et al. 2019), while our current results demonstrate that the displaced MCM is more prone to activation. Combining these two observations, we suggest that the rDNA replicates earlier in sir2 than in wild type simply because more of the MCM complex is in a location that favors its activation; in other words, we favor the view that the region of relatively low nucleosome occupancy to the right of MCM’s normal location is more amenable to MCM activation, regardless of SIR2 status. In support of this notion, close analysis of our MCM- ChEC data reveals that, while MCM complexes at the displaced location are much rarer in wild type than they are in sir2, this particular subpopulation appears equally susceptible to activation in both genotypes. This has the important implication that the manner in which the rDNA origin fires in sir2 mutants may be simply an exaggerated instance of its normal behavior, and that even in wild type cells, it is predominantly from MCMs that are displaced, likely due to low level of C- pro transcription, that replication initiates.

Classic studies that employed EM and two-dimensional gel electrophoresis showed that only a fraction of rDNA origins, estimated at one in five, are activated in any single cell cycle in yeast (Brewer and Fangman 1988; Linskens and Huberman 1988). These estimates were later corroborated by single molecule analysis, which also observed that the fraction of active origins increases in sir2 mutants (Pasero et al. 2002). This low fraction of active rDNA origins could be due to insufficient origins licensing, suppression of activation or licensed origins or a combination of the two mechanisms. Our measurements of rDNA licensing showed that MCM are loaded in up to 90% of origins in WT cells and allowed us to conclude that regulation or rDNA replication initiation occurs primarily at the level of origin firing.

Chromatin remodeling enzymes have been proposed to assist DNA replication by modulation of nucleosome phasing in the vicinity of replication origins, but it has proven difficult to assign specific roles to remodelers in genome replication in vivo (Cutler et al. 2018). One of the reasons for this difficulty is the redundancy among different remodelers on one hand, and the multiple essential roles CRE play in DNA transactions in addition to replication. For example, in vitro studies have identified four different remodelers that in combination with ORC can phase nucleosome at 320 replication origins (Chacin et al. 2023). However, a deletion of any single of these remodelers had no discernable effect on nucleosome phasing in vivo and a quadruple mutant was required to abolish phasing. While the characterization of this quadruple mutant suggested impaired replication initiation, it was difficult to exclude indirect effects that might have been caused by the effects on transcription. In contrast, the effect of a chromatin remodeler on DNA replication we describe in this work is remarkably specific: we found that the activity of a single remodeler, Fun30, is required to activation not only of a specific origin, but that its action is specific to nucleosomes adjacent to repositioned MCM complexes at the rDNA. The suppression of rDNA origin activity in sir2 mutants, which contained a large fraction of repositioned MCMs, by FUN30 deletion is so severe that it reduces the rDNA size ∼3-5 fold, from 160 to 35-50 copies.

In conclusion, our findings contribute to a growing body of evidence highlighting the importance of positioning of the MCM replicative helicase within the local nucleosome environment in determining replication timing (Azmi et al. 2017; Rodriguez et al. 2017). While it is clear that small differences in this positioning can have significant effects on the timing of DNA replication, it is unlikely that a simple rule exists to directly infer origin timing solely based on the distance between the MCM complex and the nearest nucleosome. Nonetheless, our results suggest that this parameter could be a crucial piece of the puzzle in understanding why some origins fire early in S phase, while others fire late. Further research is needed to unravel the precise molecular mechanisms governing the interplay between the MCM complex and nucleosomes and how these dynamics influence replication timing across different genomic loci. Such insights will deepen our understanding of DNA replication regulation in the context of disease states associated with epigenetic drift, such as aging and cancer.

Material and methods

Yeast strain

Strains used in the study were derived from BY4741 and W303 and are provided in Suplemental Table 1. Cells were grown in standard yeast peptone 2% dextrose media except for hydroxyurea time course study which employed synthetic media.

Strains containing point mutations in FUN30 were made using a plasmid that contains CAS9 and gRNA targeting the PAM sites adjacent to the desired mutation sites in FUN30 as described. Oligos containing FUN30-targeting gRNAs were cloned into pML104. pML104 was a gift from John Wyrick (Addgene plasmid # 67638)(Laughery et al. 2015). The plasmids, which induce double strand breaks within FUN30, were co-transformed with DNA blocks containing the repair templates with the desired mutations and altered PAM sites, synthetically synthesized by IDT. FUN30 DNA was PCR amplified from the resulting transformants and sequenced to verify that the desired mutations have been introduced.

Edu-incorporating yeast strains were derived from W303 and contain four Brdu vectors that express HSV-TK and hENT1 ((Viggiani and Aparicio 2006). Plasmids containing BrdU cassete obtained from Addgene (Addgene plasmid # 71789-71792), were integrated at four different auxotrophic loci (HIS3, TRP1, LEU2, and URA3). Single integration of each Brdu cassette was confirmed via PCR using plasmid-specific primer sets as described ((Viggiani and Aparicio 2006), which assured that the strains contain exactly four copies of the BrdU cassette

Measurements of rDNA replication timing using S-seq

S-seq experiments were carried out as previously described (Foss et al. 2017). Cells were grown to log phase in YEPD media, fixed using 70% (vol/vol) ethanol, subjected to proteinase K digestion, and their DNA stained with Sytox Green as described (Foss 2001). Cells were sorted according to DNA content on a BD Biosciences FACSAria II cell sorter into G1 and S fractions. DNA from a minimum of 1 e-6 cells from each fraction was isolated using the YeaStar Genomic DNA Kit (Zymo Research). DNA was fragmented by sonication and sequenced.

MNase seq

We carried out MNase-Seq as previously described (Foss 2019). Briefly, cells grown to log phase in rich medium, Yeast Peptone Agar with 2% glucose (YEPD), from an overnight 25 mL culture were synchronized with 3 μM alpha-factor for 1.5 hrs. at 30 ° C. Arrested cells were crosslinked with 1% formaldehyde for 30 min at room temperature water bath with shaking. Formaldehyde was quenched with 125 mM glycine and cells were centrifuged at 3000 rpm for 5 min. Cells were washed twice with water and resuspended in 1.5 mL Buffer Z (1 M sorbitol, 50 mM Tris-HCl pH 7.4) with 1 mM beta-mercaptoethanol (1.1 μL of 14.3 M beta-mercaptoethanol diluted 1:10 in Buffer Z) per 25 mL culture. Cells were treated with 100 μL 20 mg/ mL zymolyase at 30 ° C for 20–30 min. Spheroplasts were centrifuged at 5000 rpm for 10 min and resuspended in 5 mL NP buffer (1 M sorbitol, 50 mM NaCl, 10 mM Tris pH 7.4, 5 mM MgCl2, 1 mM CaCl2) supplemented with 500 μM spermidine, 1 mM beta-mercaptoethanol and 0.075% NP-40. Nuclei were aliquoted in tubes with varying concentrations of micrococcal nuclease (Worthington), mixed via tube inversion, and incubated at room temperature for 20 mins. Chromatin digested with 1.9 U– 7.5 U micrococcal nuclease per 1/5th of spheroplasts from a 25 mL culture yielded appropriate mono-, di-, tri- nucleosome protected fragments for next-generation sequencing. Digestion was stopped with freshly made 5x stop buffer (5% SDS, 50 mM EDTA) and proteinase K was added (0.2 mg/ml final concentration) for an overnight incubation at 65 ° C to reverse crosslinking. DNA was extracted with phenol/chloroform and precipitated with ethanol. Micrococcal nuclease digestion was analyzed via gel electrophoresis prior to proceeding to library preparation. Sequencing libraries were prepared as described below for ChEC.

Chromatin Endogenous Cleavage (ChEC)

ChEC-seq was carried out as previously described (Foss et al. 2019; Foss et al. 2021). Briefly, MNase was activated by addition of CaCl2 to cells that were permeabilized with digitonin. DNA was extracted from cells using phenol and chloroform, precipitated using salt and ethanol and used to construct sequencing libraries without size fractionation. Cells were centrifuged at 1,500 x g for 2 mins, and washed twice in cold Buffer A (15 mM Tris pH 7.5, 80 mM KCl, 0.1 mM EGTA) without additives. Washed cells were carefully resuspended in 570 μL Buffer A with additives (0.2 mM spermidine, 0.5 mM spermine, 1 mM PMSF, ½ cOmplete ULTRA protease inhibitors tablet, Roche, per 5 mL Buffer A) and permeabilized with 0.1% digitonin in 30 ° C water bath for 5 min. Permeabilized cells were cooled at room temperature for 1 min and 1/5th of cells were transferred in a tube with freshly made 2x stop buffer (400 mM NaCl, 20 mM EDTA, 4 mM EGTA)/1% SDS solution for undigested control. Micrococcal nuclease was activated with 5.5 μL of 200 mM CaCl2 at various times (5 mins, and 10 mins) and the reaction stopped with 2x stop buffer/1% SDS. Once all time points were collected, proteinase K was added to each of the collected time points and incubated at 55 ° C water bath for 30 mins. DNA was extracted using phenol/chloroform and precipitated with ethanol. Micrococcal nuclease digestion was analyzed via gel electrophoresis prior to proceeding to library preparation. Library was prepared as previously described used using total DNA, without any fragment size selection (Foss et al. 2019; Foss et al. 2021).

Chromatin immunoprecipitation followed by sequencing (ChIP-seq)

Cells were grown to log phase in rich medium, YEPD (1% Yeast Extract, 2% Peptone, 2% glucose), synchronized with 3 μM alpha-factor for 1.5 hrs. at 30°C. G1 arrested cells were released into 200 mM hydroxyurea (HU) YEPD and grown for 1hr. Cells were crosslinked with 1% formaldehyde for 30 min at room temperature, quenched with 125 mM glycine for 5 minutes, and centrifuged at 3000 rpm for 5 minutes. Cell pellets were washed three times with TBS (20 mM Tris-HCl, pH 7.6, 150 mM NaCl). Cell pellets were flash frozen.

Chromatin Preparation

The cell pellet was thawed in 300 μL Breaking Buffer (100 mM Tris- HCl, pH 8.0, 20% Glycerol) supplemented with protease inhibitors (Pierce Protease inhibitor EDTA-free tablets [Thermo Scientific A32965] and 1 mM PMSF). 300 μL acid-washed glass beads were added to the mixture. Cells were lysed using a BioSpec Mini-Bead-beater in a cold room: 5x 30-second pulses with 1-minute on ice between rounds. The supernatant was transferred to a new tube and added 600 μL FA buffer (50 mM Hepes-KOH, pH 7.6, 150 mM NaCl, 5 mM EDTA, 1% Triton X-100, 0.1% Sodium deoxycholate) with protease inhibitors. The chromatin was sonicated with a F60 Sonic Dismembrator 10 times for 10 seconds at setting 4, with 1-minute rests on ice between rounds. The sonicated chromatin was centrifuged at 14,000 rpm for 15 minutes at 4°C. The supernatant was transferred and spun again at 14,000 rpm for 15 minutes at 4°C.

Immunoprecipitation (IP)

For IP, 20 μL of Dynabeads Protein G beads (Invitrogen 10004D) were washed three times with 500 μL PBS-T (0.8% NaCl, 0.144% Na2HPO4, 0.02% KCl, 0.024% KH2PO, 0.1% Tween 20, pH 7.1), resuspended in PBS-T, and incubated with 5 μL of FLAG M2 (Sigma F1804) antibody for 60 minutes at room temperature. The antibody-conjugated beads were washed and resuspended in 20 μL FA buffer with protease inhibitors. IP chromatin samples (1 μg) were brought up to 400 μL with FA buffer with protease inhibitors and incubated with antibody- conjugated beads for 90 minutes at room temperature. Beads were washed three times with FA buffer with protease inhibitors, twice with FA-HS buffer (50 mM Hepes-KOH, pH 7.6, 500 mM NaCl, 5 mM EDTA, 1% Triton X-100, 0.1% Sodium deoxycholate with protease inhibitors), and once with RIPA buffer (10 mM Tris-HCl, pH 8.0, 0.25 M LiCl, 0.5% NP-40, 0.5% Sodium deoxycholate, 5 mM EDTA with protease inhibitors). Input sample was generated by adding 0.1 μg chromatin in 40 μL FA buffer to 40 μL 2X Stop buffer (20 mM Tris-HCl, pH 8.0, 100 mM NaCl, 20 mM EDTA, 1% SDS, 2% Tween 20) and reversed crosslinks by overnight incubation at 65°C.

Elution and Crosslink Reversal

Bound chromatin was eluted by incubating the beads in 80 μL 2X Stop buffer at 75°C for 10 minutes. The eluate was collected, and crosslinks were reversed by overnight incubation at 65°C. The IP and Input were treated with 4 μL of 20 mg/mL RNase A and incubated at 55°C for 1 hour, followed by 4 μL of 20 mg/mL Proteinase K and incubation at 55°C for at least 3 hours. DNA was purified using a MinElute PCR Purification Kit

Paired-Ends NGS

Chromatin immunoprecipitation samples were prepared for sequencing following the protocol for the Kapa Hyper Kit. For Input samples, 2 ng were processed for End Repair. For IP samples, the entire volume was used. PCR amplification was performed with 10- 14 cycles based on DNA quantity.

Quantification of rDNA replication using the thymidine analogue EdU

Cells were grown to log phase in rich YEPD medium from a 50 mL culture and synchronized in G1 using 3 μM alpha factor for 2 hours at 30°C. Using 500 μg/mL pronase, arrested cells were released into rich media containing 200 mM hydroxyurea and 130 μM EdU (5-Ethynyl-2’- deoxyuridine). Cells were harvested in G1 and after 60 minutes released into HU. DNA was extracted using YeaStar Genomic DNA Kit (Zymo Research) and sonicated. Separate sequencing primers were ligated to each of the samples. Samples were pooled together, and subjected to click chemistry reaction in 10 mM Tris pH 7.4, 50mM NaCl, 20 μM biotin-TEG azide (Vector Laboratories), 5mM THPTA (tris-hydroxypropyltriazolylmethylamine) (Vector Laboratories), 1mM copper sulfate, and 15 mM sodium ascorbate (total volume 200 μL). Click reaction proceeded for 30 minutes at room temperature prior to DNA ethanol precipitation. Streptavidin magnetic beads (New England Biolabs) were washed 3 times with B&W buffer I (100 mM NaCl, 10 mM Tris-HCl pH 8.0, 1 mM EDTA, 0.05% Tween20, 0.5% SDS) before adding 100 μL of bead suspension to 180 μL DNA plus B&W buffer I. After a 30-minute incubation at room temperature while gently inverting tubes, bead bound DNA was washed with 200 μL B&W buffer I followed by 100 μL Stringency wash buffer (0.1% SDS, 0.1% SSC buffer) and 200 μL Wash Buffer II (100 mM NaCl, 10 mM Tris-HCl pH 8.0, 1 mM EDTA, 0.05% Tween 20). DNA was eluted off the beads by denaturation at 98°C for 5 minutes, followed by cooling at 25°C. Edu-labeled DNA was PCR amplified post-pulldown and subjected to sequencing.

Hydoxyurea ChEC time course experiment

Cells grown in synthetic media to logarithmic phase were arrested in G1 using alpha factor for 90 minutes, centrifuged, washed in water to remove alpha factor, and released into media containing 200 mM Hydroxyurea. Cells were harvested in G1 and at indicated time points following release into HU-containing media, chilled at 4 degrees and analyzed by ChEC as described above.

qRT-PCR analysis of c-pro transcript and rDNA copy number

RNA was extracted from logarithmically growing cells after spheroplasting using an RNA-Easy column. The sequences of the primers used in qRT-PCR for c-pro and PDA1 mRNA are provided in Supplemental Table 2. rDNA size was measured by qPCR using DNA that was extracted by phenol-chloroform extraction and ethanol precipitation using the primers listed in As an internal standard, qPCR was done in parallel on the DNA extracted from fob1 strains with 180 and 35 rDNA copies whose rDNA sizes had been verified by Pulsed Field Gel Electrophoresis (Kwan et al. 2023).

ChEC Southern blots for measuring origin licensing

Cells carrying MCM2-Mnase were subjected to ChEC as described above, except that the stop SDS was omitted from the stop buffer to allow subsequent spheroplasting. DNA was extracted from spheroplasts using the YeaStar Genomic DNA Kit (Zymo Research), digested with XmnI, separated on 1.75% agarose gels and analyzed by standard Southern blotting technique. The primers used to make the rDNA and PIK1 control probe are listed in Supplemental Table 2.